<!DOCTYPE html>


<html xmlns="http://www.w3.org/1999/xhtml">
  <head>
    <meta http-equiv="Content-Type" content="text/html; charset=utf-8" />
    
    <title>Algorithms for Working with List Items - Order Code (Mixed Up Code) &mdash; Experiment 1</title>
    
    <link rel="stylesheet" href="_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="_static/codemirror.css" type="text/css" />
    <link rel="stylesheet" href="_static/activecode.css" type="text/css" />
    <link rel="stylesheet" href="_static/clickable.css" type="text/css" />
    <link rel="stylesheet" href="_static/pytutor.css" type="text/css" />
    <link rel="stylesheet" href="_static/modal-basic.css" type="text/css" />
    <link rel="stylesheet" href="_static/datafile.css" type="text/css" />
    <link rel="stylesheet" href="_static/dragndrop.css" type="text/css" />
    <link rel="stylesheet" href="_static/fitb.css" type="text/css" />
    <link rel="stylesheet" href="_static/codemirror.css" type="text/css" />
    <link rel="stylesheet" href="_static/livecode.css" type="text/css" />
    <link rel="stylesheet" href="_static/parsons.css" type="text/css" />
    <link rel="stylesheet" href="_static/lib/prettify.css" type="text/css" />
    <link rel="stylesheet" href="_static/poll.css" type="text/css" />
    <link rel="stylesheet" href="_static/tabbedstuff.css" type="text/css" />
    <link rel="stylesheet" href="_static/video.css" type="text/css" />
    <link rel="stylesheet" href="_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    
    <script type="text/javascript">
      var DOCUMENTATION_OPTIONS = {
        URL_ROOT:    './',
        VERSION:     '0.0',
        COLLAPSE_INDEX: false,
        FILE_SUFFIX: '.html',
        HAS_SOURCE:  true
      };
    </script>
    <script type="text/javascript" src="_static/runestonebase.js"></script>
    <script type="text/javascript" src="_static/jquery.js"></script>
    <script type="text/javascript" src="_static/underscore.js"></script>
    <script type="text/javascript" src="_static/doctools.js"></script>
    <script type="text/javascript" src="_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="_static/bookfuncs.js"></script>
    <script type="text/javascript" src="_static/codemirror.js"></script>
    <script type="text/javascript" src="_static/xml.js"></script>
    <script type="text/javascript" src="_static/css.js"></script>
    <script type="text/javascript" src="_static/htmlmixed.js"></script>
    <script type="text/javascript" src="_static/python.js"></script>
    <script type="text/javascript" src="_static/javascript.js"></script>
    <script type="text/javascript" src="_static/activecode.js"></script>
    <script type="text/javascript" src="_static/skulpt.min.js"></script>
    <script type="text/javascript" src="_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="_static/clike.js"></script>
    <script type="text/javascript" src="_static/timed_activecode.js"></script>
    <script type="text/javascript" src="_static/animationbase.js"></script>
    <script type="text/javascript" src="_static/mchoice.js"></script>
    <script type="text/javascript" src="_static/timedmc.js"></script>
    <script type="text/javascript" src="_static/timed.js"></script>
    <script type="text/javascript" src="_static/clickable.js"></script>
    <script type="text/javascript" src="_static/timedclickable.js"></script>
    <script type="text/javascript" src="_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="_static/pytutor.js"></script>
    <script type="text/javascript" src="_static/codelens.js"></script>
    <script type="text/javascript" src="_static/skulpt.min.js"></script>
    <script type="text/javascript" src="_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="_static/datafile.js"></script>
    <script type="text/javascript" src="_static/dragndrop.js"></script>
    <script type="text/javascript" src="_static/timeddnd.js"></script>
    <script type="text/javascript" src="_static/fitb.js"></script>
    <script type="text/javascript" src="_static/timedfitb.js"></script>
    <script type="text/javascript" src="_static/livecode.js"></script>
    <script type="text/javascript" src="_static/clike.js"></script>
    <script type="text/javascript" src="_static/lib/prettify.js"></script>
    <script type="text/javascript" src="_static/lib/hammer.min.js"></script>
    <script type="text/javascript" src="_static/parsons.js"></script>
    <script type="text/javascript" src="_static/timedparsons.js"></script>
    <script type="text/javascript" src="_static/poll.js"></script>
    <script type="text/javascript" src="_static/reveal.js"></script>
    <script type="text/javascript" src="_static/shortanswer.js"></script>
    <script type="text/javascript" src="_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="https://cdn.mathjax.org/mathjax/latest/MathJax.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script type="text/javascript" src="_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="_static/jquery-fix.js"></script>
    <script type="text/javascript" src="_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="_static/waypoints.min.js"></script>
    <script type="text/javascript" src="_static/rangy-core.js"></script>
    <script type="text/javascript" src="_static/rangy-textrange.js"></script>
    <script type="text/javascript" src="_static/rangy-cssclassapplier.js"></script>
    <script type="text/javascript" src="_static/user-highlights.js"></script>
    <script type="text/javascript" src="_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="top" title="Experiment 1" href="index.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'Ex1Parsons';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = true;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = true;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = '';
</script>

<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>



  </head>
  <body role="document">


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        <a class="navbar-brand" href="index.html">Examples Plus Practice</a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">

      <ul class="nav navbar-nav navbar-right">

        <li class="divider-vertical"></li>

        <!-- social media dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-share-alt" style="opacity: 0.9"></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li>
        <!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='index.html'>Table of Contents</a></li>
            
            <li><a href='/runestone/static/Ex1Parsons/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='http://127.0.0.1:8000/runestone/mygroup/initiateGroup' id="joinGroupLink">Join a Study Group</a></li>
            <li><a href='http://127.0.0.1:8000/runestone/mygroup/schedule' id="groupScheduleLink">Group Schedule</a></li>
            <li><a href='http://127.0.0.1:8000/runestone/mygroup/newschedule' id="newChapterLink">Schedule New Chapter</a></li>
            <li><a href='http://127.0.0.1:8000/runestone/mygroup/manageGroup' id="manageGroupLink">Manage Group</a></li>
			<li class="divider"></li>
            <li><a href='http://127.0.0.1:8000/runestone/admin/index'>Instructor's Page</a></li>
            <li><a href='http://127.0.0.1:8000/runestone/assignments/index'>Progress Page</a></li>
            <li class="divider"></li>
            <li><a href="http://127.0.0.1:8000/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="http://127.0.0.1:8000/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="http://127.0.0.1:8000/runestone/default/user/register" id="registerlink">Register</a></li>
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
          </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>

        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='/runestone/static/Ex1Parsons/navhelp.html'>Navigation Help</a></li>
            <li><a href='/runestone/static/overview/instructor.html'>Help for Instructors</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/static/help/bugreports.html'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>

      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">Algorithms for Working with List Items - Order Code (Mixed Up Code)</a><ul>
<li><a class="reference internal" href="#example-count-target-values">Example: Count Target Values</a></li>
<li><a class="reference internal" href="#practice-count-target-values-in-range">Practice: Count Target Values in Range</a></li>
<li><a class="reference internal" href="#example-find-the-maximum-value">Example: Find the Maximum Value</a></li>
<li><a class="reference internal" href="#practice-find-the-smallest-value">Practice: Find the Smallest Value</a></li>
<li><a class="reference internal" href="#example-find-the-average-of-a-list-of-values">Example: Find the Average of a List of Values</a></li>
<li><a class="reference internal" href="#practice-find-the-average-of-every-other-value-in-a-list">Practice: Find the Average of Every Other Value in a List</a></li>
<li><a class="reference internal" href="#practice-find-the-average-of-a-list-of-values-without-the-minimum-value">Practice: Find the Average of a List of Values Without the Minimum Value</a></li>
</ul>
</li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="algorithms-for-working-with-list-items-order-code-mixed-up-code">
<h1>Algorithms for Working with List Items - Order Code (Mixed Up Code)<a class="headerlink" href="#algorithms-for-working-with-list-items-order-code-mixed-up-code" title="Permalink to this headline">¶</a></h1>
<div class="section" id="example-count-target-values">
<h2>Example: Count Target Values<a class="headerlink" href="#example-count-target-values" title="Permalink to this headline">¶</a></h2>
<p>One of the common things to do with a list is to count the number of times a target value appears in a list.  To do this initialize a count to zero and then loop through all the values in the list.  If the current value is equal to the target value increment the count.  Return the count.</p>
<p>The code below also contains tests to make sure that the code works as intended.  For example <code class="docutils literal"><span class="pre">countTargetValues(5,</span> <span class="pre">[1,</span> <span class="pre">2,</span> <span class="pre">3])</span></code> should return 0 since there are no 5&#8217;s in the list and <code class="docutils literal"><span class="pre">countTargetValues(5,</span> <span class="pre">[5,</span> <span class="pre">4,</span> <span class="pre">5])</span></code> should return 2 since there are two 5&#8217;s in that list.</p>

<div data-childcomponent="Count_Targets" class="explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=Count_Targets data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
# define the function
def countTargetValues(target, numList):

    # initialize the count
    count = 0

    # loop through all the indices
    for index in range(len(numList)):

        # get the current value
        current = numList[index]

        # if the curent value is equal to the given target value
        if (current == target):

            # increment the count
            count = count + 1

    # return the count
    return count

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testTarget(self):
        self.assertEqual(countTargetValues(5, [1, 2, 3]), 0)
        self.assertEqual(countTargetValues(5, [3, 2, 5]), 1)
        self.assertEqual(countTargetValues(5, [5, 4, 5]), 2)
        self.assertEqual(countTargetValues(5, [5, 5, 5]), 3)

myTests().main()
</textarea>
</div>
</div>
<div class="section" id="practice-count-target-values-in-range">
<h2>Practice: Count Target Values in Range<a class="headerlink" href="#practice-count-target-values-in-range" title="Permalink to this headline">¶</a></h2>
<p>You can also count the number of target items in a list between a first and last index, including the values at the first and last indices.  To do this initialize a count to zero and then loop through all the values in the list from the start index to the end index. If the current value is equal to the target value increment the count. Return the count.</p>
<div class="admonition note">
<p class="first admonition-title">Note</p>
<p class="last">Remember that the <code class="docutils literal"><span class="pre">range(first,</span> <span class="pre">last)</span></code> function actually returns the values from first to last - 1.  So <code class="docutils literal"><span class="pre">range(1,3)</span></code> returns <code class="docutils literal"><span class="pre">[1,2]</span></code>.</p>
</div>
<p>The following function should count and return the number of values in a list that are equal to a target value in the range specified by the start and end indicies (including the values at the start and end indicies).</p>
<p>For example <code class="docutils literal"><span class="pre">countTargetInRange(3,</span> <span class="pre">0,</span> <span class="pre">1,</span> <span class="pre">[3,</span> <span class="pre">3,</span> <span class="pre">2,</span> <span class="pre">6])</span></code> should return 2 since there are 2 values that are greater than or equal to 3 from index 0 to 1.  The call <code class="docutils literal"><span class="pre">countTargetInRange(3,</span> <span class="pre">0,</span> <span class="pre">3,</span> <span class="pre">[2,</span> <span class="pre">-3,</span> <span class="pre">0,</span> <span class="pre">1])</span></code> should return 0 since there are no values that are greater than or equal to 3 in that list.</p>

    <pre data-component="parsons" id="Count_Target_In_Range1">
        <span data-question>2-1-1: The code below is mixed up and contains extra blocks that are not needed.  Drag the needed code from the left to the right and put them in order with the correct indention so that the code would work correctly.</span>def countInRange(target, start, end, numList):
---
def countInRange(target, start, end, numList) #paired
---
    count = 0
---
    count = 1 #paired
---
    for index in range(start, end+1):
---
        current = numList[index]
---
        if current == target:
---
            count = count + 1
---
            count = Count + 1 #paired
---
    return count
---
    return coun #paired</pre>
    </div>
<div class="section" id="example-find-the-maximum-value">
<h2>Example: Find the Maximum Value<a class="headerlink" href="#example-find-the-maximum-value" title="Permalink to this headline">¶</a></h2>
<p>Another common thing to do is find the maximum value in a list.  To do this create a variable to keep track of the maximum value found so far and set it to the first item in the list to start.  Then loop through all the values in the list and each time test if the current value is greater than the saved maximum value.  If it is, then set the maximum value to the current value.  Return the maximum value.</p>
<p>The code below also contains tests to make sure that the code works as intended. For example <code class="docutils literal"><span class="pre">getMaximum([9,</span> <span class="pre">3,</span> <span class="pre">8])</span></code> should return 9 and <code class="docutils literal"><span class="pre">getMaximum([-2,</span> <span class="pre">-1,</span> <span class="pre">-3]),</span> <span class="pre">-1)</span></code> should return -1.</p>

<div data-childcomponent="Get_Max" class="explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=Get_Max data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
# define the function
def getMaximum(numList):

    # set the maximum to the first item
    maximum = numList[0]

    # loop through all the values
    for current in numList:

        # if num is greater than the maximum so far
        if current > maximum:

            # set the maximum to the current value
            maximum = current

    # return the maximum value
    return maximum

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testTarget(self):
        self.assertEqual(getMaximum([9, 3, 8]), 9)
        self.assertEqual(getMaximum([-2, -1, -3]), -1)
        self.assertEqual(getMaximum([2, 3, 5, 15]), 15)
        self.assertEqual(getMaximum([32, 64, 28]), 64)
        self.assertEqual(getMaximum([3]), 3)

myTests().main()
</textarea>
</div>
</div>
<div class="section" id="practice-find-the-smallest-value">
<h2>Practice: Find the Smallest Value<a class="headerlink" href="#practice-find-the-smallest-value" title="Permalink to this headline">¶</a></h2>
<p>The following function should find and return the minimum in a list.  To find the minimum, create a variable to keep track of the minimum value found so far and set it to the first item in the list to start. Then loop through all the values in the list and each time test if the current value is less than the saved smallest value. If it is, then set the smallest value to the current value. Return the smallest value.</p>
<p>For example <code class="docutils literal"><span class="pre">getMinimum([9,</span> <span class="pre">3,</span> <span class="pre">8])</span></code> should return 3 and <code class="docutils literal"><span class="pre">getMinimum([-2,</span> <span class="pre">-1,</span> <span class="pre">-3])</span></code> should return -3.</p>

    <pre data-component="parsons" id="Find_Min1">
        <span data-question>2-1-2: The code below is mixed up and contains extra blocks that are not needed.  Drag the needed code from the left to the right and put them in order with the correct indention so that the code would work correctly.</span>def getMinimum(numList):
---
def getMinimum(numList: #paired
---
    minimum = numList[0]
---
    for current in numList:
---
    for current in numlist: #paired
---
        if current < minimum:
---
        if current > minimum #paired
---
            minimum = current
---
            min = current #paired
---
    return minimum</pre>
    </div>
<div class="section" id="example-find-the-average-of-a-list-of-values">
<h2>Example: Find the Average of a List of Values<a class="headerlink" href="#example-find-the-average-of-a-list-of-values" title="Permalink to this headline">¶</a></h2>
<p>Another common thing to do to a list of values is to find the average of them.   To calculate an average first create a variable to hold the sum and set it to zero.  Then loop through all the values in the list and add each one to the sum.  Remember that you can&#8217;t divide by 0, so if the list is empty return 0 and otherwise return the sum divided by the number of items in the list.</p>
<p>The code below also contains tests to make sure that the code works as intended. For example <code class="docutils literal"><span class="pre">getAverage([50,60,70])</span></code> should return 60 and <code class="docutils literal"><span class="pre">getAverage([])</span></code> should return 0.</p>

<div data-childcomponent="Get_Average" class="explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=Get_Average data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
def getAverage(numList):

    # initialize the sum
    sum = 0

    # loop through the values
    for num in numList:

        # add the value to the sum
        sum = sum + num

    # prevent a divide by zero if this is called with an empty list
    if len(numList) == 0:
        return 0
    else:
        return sum / len(numList)

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(getAverage([50,60,70]),60)
        self.assertEqual(getAverage([75,60,80,95]),77.5)
        self.assertEqual(getAverage([10,20,30,40,90]),38)
        self.assertEqual(getAverage([]),0)
        self.assertEqual(getAverage([4]),4)

myTests().main()
</textarea>
</div>
</div>
<div class="section" id="practice-find-the-average-of-every-other-value-in-a-list">
<h2>Practice: Find the Average of Every Other Value in a List<a class="headerlink" href="#practice-find-the-average-of-every-other-value-in-a-list" title="Permalink to this headline">¶</a></h2>
<p>The following code should find the average of every other value in a list starting with the value at index 0.  To do this first initialize a sum and count to zero. Next loop through every other index starting at index zero, add the value at each index to the sum, and increment the count.  If the count is zero then return zero and otherwise return the sum divided by the count.</p>
<p>The code below contains code for testing the <code class="docutils literal"><span class="pre">getAverageEveryOther</span></code> function.  For example <code class="docutils literal"><span class="pre">getAverageEveryOther([10,30,90])</span></code> should return 50 since 10 plus 90 is 100 and 100 divided by 2 is 50, and <code class="docutils literal"><span class="pre">getAverageEveryOther([])</span></code> should return 0 since the count would be zero.</p>

    <pre data-component="parsons" id="Get_Avg_Every_Other1">
        <span data-question>2-1-3: The code is mixed up and contains extra blocks that are not needed.  Drag the needed code from the left to the right and put them in order with the correct indention so that the code would work correctly.</span>def getAverageEveryOther(numList):
---
    sum = 0
    count = 0
---
    for index in range(0,len(numList),2):
---
    for index in range(0,numList,2): #paired
---
        value = numList[index]
        sum = sum + value
        count = count + 1
---
        value = numList[index]
        sum = value
        count = count + 1 #paired
---
    if count == 0:
---
        return 0
---
    else:
---
    else #paired
---
        return sum / count
---
        sum / count #paired</pre>
    </div>
<div class="section" id="practice-find-the-average-of-a-list-of-values-without-the-minimum-value">
<h2>Practice: Find the Average of a List of Values Without the Minimum Value<a class="headerlink" href="#practice-find-the-average-of-a-list-of-values-without-the-minimum-value" title="Permalink to this headline">¶</a></h2>
<p>The following code should find the average of a list of values except it should not include any of the minimum value in the list in the calculation of the average. It should keep track of the number of values used in the sum and return the average as the sum divided by the count of values used in the sum. If there are no values in the sum (the count of values used in the sum is 0) then it should return 0.</p>
<p>For example, <code class="docutils literal"><span class="pre">getAverageExceptMinimum[90,</span> <span class="pre">3,</span> <span class="pre">10,</span> <span class="pre">3])</span></code> should return 50 since 90 plus 10 is 100 (ignoring both 3&#8217;s) and the average is 100 divided by 2 which is 50.  And <code class="docutils literal"><span class="pre">getAverageExceptMinimum([2])</span></code> should return 0 since removing the smallest value means the number of values in the sum is 0 so the average should be 0.</p>

    <pre data-component="parsons" id="Get_Avg_Except_Minimum1">
        <span data-question>2-1-4: The code is mixed up and contains extra blocks that are not needed.  Drag the needed code from the left to the right and put them in order with the correct indention so that the code would work correctly.</span>def getAverageExceptMinimum(numList):
---
    min = numList[0]
---
    min = numList[0 #paired
---
    for num in numList:
---
        if num < min:
---
            min = num
---
            num = min #paired
---
    sum = 0
    count = 0
---
    sum = 0
    count = 1 #paired
---
    for num in numList:
---
        if num != min:
---
            sum = sum + num
            count = count + 1
---
    if count == 0:
---
    if sum == 0: #paired
---
        return 0
---
    else:
---
        return sum / count</pre>
    </div>
</div>


  
      

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      <span id='numuserspan'></span> readers online now | <span class='loggedinuser'></span>
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2015 CS Learning 4 U group at Georgia Tech.
      Created using <a href="http://sphinx.pocoo.org/">Sphinx</a> 1.4.5.
    </p>
  </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>